<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>Interactive: Point Placement Graph</text></name>
    <questiontext format="html">
      <text><![CDATA[Sketch the function \(f(t)\) by clicking on the graph to place points at the key locations.<br><small>Click to place 5 points. They will connect automatically.</small>

<div class="stack-part"><p><strong>(a)</strong> Place 5 points on the graph:</p><div class="jsxgraph-box">
[[jsxgraph input-ref-ans1="ans1Ref" width="500px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-5, 70, 65, -70],
    axis: true, showNavigation: true, showCopyright: false, grid: true
});

board.create('text', [62, -5, 't (s)'], {fontSize: 14});
board.create('text', [-3, 65, 'f(t)'], {fontSize: 14});

var studentPoints = [];
var segments = [];
var stateInput = document.getElementById(ans1Ref);

function serializePoints() {
    if (studentPoints.length === 0) {
        stateInput.value = '[]';
    } else {
        var parts = [];
        for (var i = 0; i < studentPoints.length; i++) {
            parts.push('[' + Math.round(studentPoints[i].X()) + ',' + Math.round(studentPoints[i].Y()) + ']');
        }
        stateInput.value = '[' + parts.join(',') + ']';
    }
    stateInput.dispatchEvent(new Event('change'));
}

function getMouseCoords(e) {
    var cPos = board.getCoordsTopLeftCorner(e),
        absPos = JXG.getPosition(e),
        dx = absPos[0] - cPos[0],
        dy = absPos[1] - cPos[1];
    return new JXG.Coords(JXG.COORDS_BY_SCREEN, [dx, dy], board);
}

function addPoint(x, y) {
    var p = board.create('point', [x, y], {
        name: '(' + Math.round(x) + ',' + Math.round(y) + ')',
        size: 4, face: 'o', strokeColor: '#2563eb', fillColor: '#2563eb',
        snapToGrid: true
    });
    p.on('drag', function() {
        p.setName('(' + Math.round(p.X()) + ',' + Math.round(p.Y()) + ')');
        board.update();
        serializePoints();
    });
    studentPoints.push(p);
    if (studentPoints.length > 1) {
        segments.push(board.create('segment',
            [studentPoints[studentPoints.length-2], studentPoints[studentPoints.length-1]],
            {strokeColor: '#ef4444', strokeWidth: 2}
        ));
    }
    return p;
}

/* Restore state from saved answer on page reload */
if (stateInput.value && stateInput.value !== '' && stateInput.value !== '[]') {
    var matches = stateInput.value.match(/\[(-?[\d.]+),\s*(-?[\d.]+)\]/g);
    if (matches) {
        for (var i = 0; i < matches.length && i < 5; i++) {
            var nums = matches[i].match(/(-?[\d.]+)/g);
            if (nums && nums.length >= 2) addPoint(parseFloat(nums[0]), parseFloat(nums[1]));
        }
        board.update();
    }
}

board.on('down', function(e) {
    var canCreate = true, coords, el, x, y;

    if (e[JXG.touchProperty]) return;

    coords = getMouseCoords(e);

    for (el in board.objects) {
        if (JXG.isPoint(board.objects[el]) && board.objects[el].hasPoint(coords.scrCoords[1], coords.scrCoords[2])) {
            canCreate = false;
            break;
        }
    }

    if (!canCreate || studentPoints.length >= 5) return;

    x = Math.round(coords.usrCoords[1]);
    y = Math.round(coords.usrCoords[2]);
    addPoint(x, y);
    serializePoints();
});

board.create('button', [5, 60, 'Reset', function() {
    for (var i = segments.length - 1; i >= 0; i--) board.removeObject(segments[i]);
    for (var i = studentPoints.length - 1; i >= 0; i--) board.removeObject(studentPoints[i]);
    studentPoints = [];
    segments = [];
    board.update();
    serializePoints();
}]);
[[/jsxgraph]]
</div>
<p style="display:none">[[input:ans1]] [[validation:ans1]]</p></div>]]></text>

    </questiontext>

    <generalfeedback format="html">
      <text><![CDATA[The function has key points at t=0, t={@t1@}, t={@t2@}, t={@t3@}, and t={@t4@}.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>

    <hint format="html">
      <text><![CDATA[Start from the origin (0,0) and work through each time interval.]]></text>
    </hint>
    <stackversion><text>2025040100</text></stackversion>
    <questionvariables>
      <text><![CDATA[t1: 10;
t2: 30;
t3: 40;
t4: 60;
p0: [0, 0];
p1: [t1, 5*t1];
p2: [t2, -5*t2 + 100];
p3: [t3, -50];
p4: [t4, 0];
correct_points: [p0, p1, p2, p3, p4];
ta1: correct_points;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[[[feedback:prt1]]]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[t1={@t1@}, t2={@t2@}, t3={@t3@}, t4={@t4@}, ans1={@ta1@}]]></text>
    </questionnote>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>0</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      
      <feedbackvariables>
        <text><![CDATA[/* Convert student answer from matrix to list (Maxima parses [[a,b]] as matrix) */
student_raw: ans1;
student_pts: if matrixp(student_raw) then args(student_raw) else student_raw;

/* Check number of points */
student_count: length(student_pts);
correct_count: is(student_count = 5);

/* Tolerance for point matching */
tolerance: 5;

/* Nearest-point matching (order-independent) */
matched_expected: makelist(false, i, 1, 5);
matched_student: makelist(false, i, 1, student_count);
match_pair: makelist(0, i, 1, 5);

match_results: block([],
    for ei: 1 thru 5 do (
        block([best_dist, best_si, ex, ey, d],
            ex: correct_points[ei][1],
            ey: correct_points[ei][2],
            best_dist: 99999,
            best_si: 0,
            for si: 1 thru student_count do (
                if not matched_student[si] then (
                    d: sqrt((student_pts[si][1] - ex)^2 + (student_pts[si][2] - ey)^2),
                    if is(d < best_dist) then (best_dist: d, best_si: si)
                )
            ),
            if is(best_si > 0) and is(best_dist < tolerance) then (
                matched_expected[ei]: true,
                matched_student[best_si]: true,
                match_pair[ei]: best_si
            )
        )
    ),
    matched_expected
);

pt_checks: makelist(if matched_expected[i] then 1 else 0, i, 1, 5);
num_correct: apply("+", pt_checks);
all_correct: correct_count and is(num_correct = 5);

/* Build detailed feedback table showing which points matched */
feedback_parts: makelist(
    sconcat(
        "<tr><td>", i, "</td>",
        "<td>(", string(correct_points[i][1]), ", ", string(correct_points[i][2]), ")</td>",
        "<td>",
        if is(match_pair[i] > 0) then
            sconcat("(", string(student_pts[match_pair[i]][1]), ", ", string(student_pts[match_pair[i]][2]), ")")
        else "no match",
        "</td>",
        "<td>", if is(pt_checks[i] = 1) then "&#10004;" else "&#10008;", "</td></tr>"
    ),
    i, 1, 5
);

feedback_msg: sconcat(
    "<strong>", string(num_correct), " / 5 points correct</strong>",
    "<table border='1' cellpadding='4' style='border-collapse:collapse;margin-top:8px'>",
    "<tr><th>#</th><th>Expected</th><th>Your closest</th><th>Result</th></tr>",
    simplode(feedback_parts, ""),
    "</table>"
);
/* Ensure feedback_msg exists */
if not boundp(feedback_msg) then feedback_msg: "";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <answertest>AlgEquiv</answertest>
        <sans>all_correct</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html"><text><![CDATA[Correct! All points are in the right positions.<br><hr><p>{@feedback_msg@}</p>]]></text></truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html"><text><![CDATA[Some points are not in the correct positions. Check each key point of the function.<br><hr><p>{@feedback_msg@}</p>]]></text></falsefeedback>
      </node>
    </prt>
    <deployedseed>12345</deployedseed>
    <deployedseed>10101</deployedseed>
    <deployedseed>10102</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Model answer earns full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1, simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>